// Copyright (C) 2020 Mozilla Corporation. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: pending
description: >
  The Iterator constructor is a built-in function
features: [iterator-helpers]
---*/

assert.sameValue(typeof Iterator, 'function');
