// Copyright 2020 Mozilla Corporation. All rights reserved.
// This code is governed by the license found in the LICENSE file.
/*---
esid: pending
description: The AsyncIterator constructor is a built-in function.
features: [iterator-helpers]
---*/

assert.sameValue(typeof AsyncIterator, 'function');



